# Loops

Topics:
1. `while` loops
2. `for` loops
<br>

## while loops

The syntax for a `while` is

```julia
while *condition*
    *loop body*
end
```

For example, we could use `while` to count or to iterate over an array.

```julia
n = 0
while n < 10
    n += 1
    println(n)
end
n
```

Output:
```
1
2
3
4
5
6
7
8
9
10
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

i = 1
while i <= length(myfriends)
    friend = myfriends[i]
    println("Hi $friend, it's great to see you!")
    i += 1
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

## for loops

The syntax for a `for` loop is

```julia
for *var* in *loop iterable*
    *loop body*
end
```

We could use a for loop to generate the same results as either of the examples above:

```julia
for n in 1:10
    println(n)
end
```

Output:
```
1
2
3
4
5
6
7
8
9
10
```

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]

for friend in myfriends
    println("Hi $friend, it's great to see you!")
end
```

Output:
```
Hi Ted, it's great to see you!
Hi Robyn, it's great to see you!
Hi Barney, it's great to see you!
Hi Lily, it's great to see you!
Hi Marshall, it's great to see you!
```

Now let's use `for` loops to create some addition tables, where the value of every entry is the sum of its row and column indices. <br>

Note that we iterate over this array via column-major loops in order to get the best performance. More information about fast indexing of multidimensional arrays inside nested loops can be found at https://docs.julialang.org/en/v1/manual/performance-tips/#Access-arrays-in-memory-order,-along-columns-1

First, we initialize an array with zeros.

```julia
m, n = 4, 5
A = fill(0, (m, n))
```

Output:
```
4×5 Matrix{Int64}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
for j in 1:n
    for i in 1:m
        A[i, j] = i + j
    end
end
A
```

Output:
```
4×5 Matrix{Int64}:
 2  3  4  5  6
 3  4  5  6  7
 4  5  6  7  8
 5  6  7  8  9
```

Here's some syntactic sugar for the same nested `for` loop

```julia
B = fill(0, (m, n))
```

Output:
```
5×5 Matrix{Int64}:
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
 0  0  0  0  0
```

```julia
for j in 1:n, i in 1:m
    B[i, j] = i + j
end
B
```

Output:
```
5×5 Matrix{Int64}:
 2  3  4  5   6
 3  4  5  6   7
 4  5  6  7   8
 5  6  7  8   9
 6  7  8  9  10
```

The more "Julia" way to create this addition table would have been with an *array comprehension*.

```julia
C = [i + j for i in 1:m, j in 1:n]
```

Output:
```
4×5 Matrix{Int64}:
 2  3  4  5  6
 3  4  5  6  7
 4  5  6  7  8
 5  6  7  8  9
```

### Exercises

#### 4.1 
Loop over integers between 1 and 100 and print their squares.

```julia
for i in 1:100
    println("i = $i, i^2 = $(i^2)")
    end
```

Output:
```
i = 1, i^2 = 1
i = 2, i^2 = 4
i = 3, i^2 = 9
i = 4, i^2 = 16
i = 5, i^2 = 25
i = 6, i^2 = 36
i = 7, i^2 = 49
i = 8, i^2 = 64
i = 9, i^2 = 81
i = 10, i^2 = 100
i = 11, i^2 = 121
i = 12, i^2 = 144
i = 13, i^2 = 169
i = 14, i^2 = 196
i = 15, i^2 = 225
i = 16, i^2 = 256
i = 17, i^2 = 289
i = 18, i^2 = 324
i = 19, i^2 = 361
i = 20, i^2 = 400
i = 21, i^2 = 441
i = 22, i^2 = 484
i = 23, i^2 = 529
i = 24, i^2 = 576
i = 25, i^2 = 625
i = 26, i^2 = 676
i = 27, i^2 = 729
i = 28, i^2 = 784
i = 29, i^2 = 841
i = 30, i^2 = 900
i = 31, i^2 = 961
i = 32, i^2 = 1024
i = 33, i^2 = 1089
i = 34, i^2 = 1156
i = 35, i^2 = 1225
i = 36, i^2 = 1296
i = 37, i^2 = 1369
i = 38, i^2 = 1444
i = 39, i^2 = 1521
i = 40, i^2 = 1600
i = 41, i^2 = 1681
i = 42, i^2 = 1764
i = 43, i^2 = 1849
i = 44, i^2 = 1936
i = 45, i^2 = 2025
i = 46, i^2 = 2116
i = 47, i^2 = 2209
i = 48, i^2 = 2304
i = 49, i^2 = 2401
i = 50, i^2 = 2500
i = 51, i^2 = 2601
i = 52, i^2 = 2704
i = 53, i^2 = 2809
i = 54, i^2 = 2916
i = 55, i^2 = 3025
i = 56,
```

#### 4.2 
Add to the code above a bit to create a dictionary, `squares` that holds integers and their squares as key, value pairs such that

```julia
squares[10] == 100
```

```julia
squares = Dict(i => i^2 for i in 1:100)
```

Output:
```
Dict{Int64, Int64} with 100 entries:
  5  => 25
  56 => 3136
  35 => 1225
  55 => 3025
  60 => 3600
  30 => 900
  32 => 1024
  6  => 36
  67 => 4489
  45 => 2025
  73 => 5329
  64 => 4096
  90 => 8100
  4  => 16
  13 => 169
  54 => 2916
  63 => 3969
  86 => 7396
  91 => 8281
  ⋮  => ⋮
```

<!-- grade_id: cell-54834a2e4e122b70, points: 1, grade: true, solution: false -->
```julia
@assert squares[10] == 100
@assert squares[11] == 121
```

#### 4.3 
Use an array comprehension to create an an array `squares_arr` that stores the squares for all integers between 1 and 100.

```julia
squares_arr = [i^2 for i in 1:100]
```

Output:
```
100-element Vector{Int64}:
     1
     4
     9
    16
    25
    36
    49
    64
    81
   100
     ⋮
  8464
  8649
  8836
  9025
  9216
  9409
  9604
  9801
 10000
```

<!-- grade_id: cell-17d4631c51b8b4f7, points: 1, grade: true, solution: false -->
```julia
@assert length(squares_arr) == 100
@assert sum(squares_arr) == 338350
```

```julia
function recur_factorial(n)
    if n == 0
        return 1
    else
        return n * recur_factorial(n - 1)
    end
end
@assert recur_factorial(5) == factorial(5)


```

```julia
function recur_sum(n)
    if n == 0
        return 0
    end
    a = 0
    for i in 1:n
        a = a + i
    end
    return a
end
@assert recur_sum(5) == 15

b = 1e5
println("b = $b")
```

Output:
```
b = 100000.0
```

Writing my own insertion sort!

```julia
# The code should be able to sort a small array in an ascending order from left to right. The first j-1 elements should always be sorted
# and the jth element then finds its position within them by being moved to a lefter position until the element to its left is less than it.

function insertion_sort_v1(A)
    for j in 2:length(A)
        key = A[j]
        i = j - 1

        while i > 0
            if A[i] > key  # put this condition in the while loop because otherwise the while loop keeps on 
                           # running even after the correct spot for key has been found
                A[i + 1] = A[i]
                A[i] = key # this is not needed because if key needs to be moved two spots then on the second 
                           # run of the loop you will rewrite the key assignment of the previous run through 
                           # A[i+1] = A[i]. Instead put A[i + 1] = key after the while loop.
                
            end
            i -= 1
        end
    end
    return A
end

function insertion_sort_updated(A)
    for j in 2:length(A)
        key = A[j]
        i = j - 1
        while i > 0 && A[i] > key
            A[i + 1] = A[i]
            i -= 1
        end
        A[i + 1] = key
    end
    return A
end
```

Output:
```
insertion_sort_updated (generic function with 1 method)
```

```julia
using BenchmarkTools, Random

N = 50_000
data = rand(1:1_000_000, N)

@btime begin v = copy($data); insertion_sort_v1(v); end
@btime begin v = copy($data); insertion_sort_updated(v); end
```

Output:
```
584.704 ms (3 allocations: 400.06 KiB)
  293.993 ms (3 allocations: 400.06 KiB)
50000-element Vector{Int64}:
    110
    163
    168
    196
    235
    263
    352
    371
    398
    400
      ⋮
 999858
 999874
 999882
 999950
 999960
 999971
 999982
 999984
 999992
```

```julia
@btime begin
    v = copy($data)
    sort!(v)
end
```

Output:
```
165.084 μs (9 allocations: 810.19 KiB)
50000-element Vector{Int64}:
    110
    163
    168
    196
    235
    263
    352
    371
    398
    400
      ⋮
 999858
 999874
 999882
 999950
 999960
 999971
 999982
 999984
 999992
```

```julia
# picking entries from an array uniformly at random with replacement
# generate an array of integers
arr = [i*100 for i in 1:410]
# generate indices to be picked from arr
idx = rand(1:410, 410)
println(idx)
# new array with values sampled from arr
arr_b = arr[idx]
println(arr_b)

```

Output:
```
[410, 165, 49, 20, 323, 169, 356, 37, 97, 176, 275, 206, 126, 352, 130, 174, 270, 377, 298, 340, 276, 321, 208, 29, 290, 323, 349, 188, 276, 327, 167, 184, 191, 210, 295, 292, 312, 383, 55, 340, 230, 145, 64, 223, 277, 212, 124, 238, 302, 139, 334, 302, 155, 99, 367, 144, 384, 97, 261, 282, 247, 135, 175, 41, 132, 283, 406, 351, 85, 390, 241, 20, 68, 205, 239, 189, 347, 383, 279, 353, 101, 288, 348, 350, 255, 150, 77, 177, 133, 37, 138, 302, 406, 140, 19, 398, 7, 401, 62, 210, 151, 27, 116, 269, 250, 393, 15, 261, 388, 172, 262, 389, 219, 85, 403, 125, 25, 81, 303, 331, 267, 256, 151, 332, 105, 19, 38, 251, 49, 358, 124, 217, 409, 404, 182, 118, 288, 20, 34, 37, 387, 150, 129, 24, 59, 168, 115, 309, 64, 336, 343, 20, 142, 99, 45, 237, 246, 15, 312, 193, 307, 367, 289, 12, 65, 390, 158, 81, 248, 191, 81, 397, 343, 1, 6, 136, 297, 52, 41, 175, 116, 392, 120, 217, 128, 170, 64, 289, 111, 332, 151, 303, 401, 178, 130, 108, 295, 11, 248, 131, 125, 37, 101, 218, 218, 290, 173, 156, 334, 298,
```

Please click on `Validate` on the top, once you are done with the exercises.